For this problem set, we'll be using the Jupyter notebook:

# Part A (2 points)

Write a function that returns a list of numbers, such that xi=i2
, for 1≤i≤n. Make sure it handles the case where n<1 by raising a ValueError.

<!-- grade_id: cell-5973d610255a807e, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Your function should print [1, 4, 9, 16, 25, 36, 49, 64, 81, 100] for n=10. Check that it does:

<!-- grade_id: cell-bf706308c14d2e6e, points: 2, grade: true, solution: false -->
```python
squares(10)

"""Check that squares returns the correct output for several inputs"""
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

<!-- grade_id: cell-be4b50672f134179, points: 2, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
try:
    squares(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    squares(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

# Part B (1 point)

Using your squares function, write a function that computes the sum of the squares of the numbers from 1 to n. Your function should call the squares function -- it should NOT reimplement its functionality.

<!-- grade_id: cell-051621a83f104148, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

<!-- grade_id: cell-7222a0df8724e01c, points: 0, grade: true, solution: false -->
```python
sum_of_squares(10)
```

<!-- grade_id: cell-a5b8499bf5350242, points: 2, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
assert sum_of_squares(1) == 1
assert sum_of_squares(2) == 5
assert sum_of_squares(10) == 385
assert sum_of_squares(11) == 506
```

<!-- grade_id: cell-1e198d95adab2905, points: 2, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""
orig_squares = squares
del squares
try:
    sum_of_squares(1)
except NameError:
    pass
else:
    raise AssertionError("sum_of_squares does not use squares")
finally:
    squares = orig_squares
```

# Part C (1 point)

Using LaTeX math notation, write out the equation that is implemented by your sum_of_squares function.

∑ni=1i2

# Part D (2 points)

Find a usecase for your sum_of_squares function and implement that usecase in the cell below.

<!-- grade_id: cell-b4ed903f7affadd1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```