# Part A (3 points)

Write a function that returns a list of numbers, such that xi=i2 , for 1≤i≤n. Make sure it handles the case where n<1 by raising a ValueError.

<!-- grade_id: cell-5dc080c7963fae78, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    ### BEGIN SOLUTION
    if n<1:
        raise ValueError
    data = []
    for i in range(n):
        i += 1
        data.append(i*i)
    return data
    ### END SOLUTION
```

Your function should return [1, 4, 9, 16, 25, 36, 49, 64, 81, 100] for n=10. Check that it does:

```python
# See what squares returns (you can change the value to see what the function returns)
squares(10)
```

<!-- grade_id: cell-7bcf2ae00414cbb4, points: 2, grade: true, solution: false -->
```python

"""Check that squares returns the correct output for several inputs"""
assert squares(1) == [1]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
### BEGIN HIDDEN TESTS
assert squares(2) == [1, 4]
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
### END HIDDEN TESTS
```

<!-- grade_id: cell-7a80e5aaf5da35bd, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
try:
    squares(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    squares(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

# Part B (2 point)¶

Using your squares function, write a function that computes the sum of the squares of the numbers from 1 to n. Your function should call the squares function -- it should NOT reimplement its functionality.

<!-- grade_id: cell-75cc5d21a954e786, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(m):
    """Compute the sum of the squares of numbers from 1 to n."""
    ### BEGIN SOLUTION
    return sum(squares(m))
    ### END SOLUTION
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```python
# See what sum_of_squares returns (you can change the value to see what the function returns)
sum_of_squares(10)
```

<!-- grade_id: cell-62648d319eded5cf, points: 1, grade: true, solution: false -->
```python


"""Check that sum_of_squares returns the correct answer for various inputs."""
assert sum_of_squares(1) == 1
assert sum_of_squares(10) == 385
### BEGIN HIDDEN TESTS
assert sum_of_squares(2) == 5
assert sum_of_squares(11) == 506
### END HIDDEN TESTS
```

<!-- grade_id: cell-108949b47d36371a, points: 1, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""
orig_squares = squares
del squares
try:
    sum_of_squares(1)
except NameError:
    pass
else:
    raise AssertionError("sum_of_squares does not use squares")
finally:
    squares = orig_squares
```